As this is a case study individual assignment, I agree and acknowledge that all code modified in this notebook is my own. I have not and will not collaborate with anyone on this assignment. If I have questions, I will ask the instructor or TAs.

<!-- grade_id: cell-aa7ec3e56dbc7904, points: 0, grade: false, solution: true -->
```python
# Please provide your name and agreement to the above statement as variables `name` (string) and `agree` (boolean)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-866f40538c74ecb9, points: 0, grade: true, solution: false -->
```python
conditions = [isinstance(name, str), isinstance(agree, bool), agree]
for condition in conditions:
    if not condition:
        raise ValueError("Student has not agreed to work on this assignment alone and without collaboration or has not provided their name")
```

# Python and Pandas

In this notebook of the case study, you'll go over an overview of [Pandas](https://pandas.pydata.org/). If you're already familiar with them, this should be a good review to prepare you for the rest of the class.

We'll go over the following:

1. Reading data in an incorrect format
1. `NaN`s
1. Series and DataFrames
1. One hot encoding
1. Boolean indexing
1. Applying functions to data
1. Creating new DataFrame columns
1. Grouping a DataFrame


```python
import pandas as pd
```

## Reading Data

Pandas can read data in a variety of ways. We'll be using the titanic dataset that is available as a csv file so we will use [`read_csv`](https://pandas.pydata.org/pandas-docs/stable/generated/pandas.read_csv.html), however you should check out the [other methods](https://pandas.pydata.org/pandas-docs/stable/api.html#flat-file).

```python
df = pd.read_csv("http://biostat.mc.vanderbilt.edu/wiki/pub/Main/DataSets/titanic3.csv")
```

Usually when we read a file in pandas, it will create a [DataFrame](https://pandas.pydata.org/pandas-docs/stable/generated/pandas.DataFrame.html) object. You can read more about dataframes in the [pandas guide](https://pandas.pydata.org/pandas-docs/stable/dsintro.html#dataframe) .You can preview the top and bottom of the data using the [head](https://pandas.pydata.org/pandas-docs/stable/generated/pandas.DataFrame.head.html) and [tail](https://pandas.pydata.org/pandas-docs/stable/generated/pandas.DataFrame.tail.html) methods of a pandas DataFrame.

```python
df.head()
```

<!-- grade_id: cell-c285b9edfcfbb9c8, points: 10, grade: true, solution: true -->
```python
# Try out the tail method here
# YOUR CODE HERE
raise NotImplementedError()
```

You can read more about the data on the [main data page](http://biostat.mc.vanderbilt.edu/wiki/pub/Main/DataSets/titanic.html) and the [version 3 data description](http://biostat.mc.vanderbilt.edu/wiki/pub/Main/DataSets/Ctitanic3.html).

As a summary, here's what the column names for the data actually represent:

- `pclass` - passenger class - 1st, 2nd or 3rd class
- `survived` - whether or not the passenger survived. 1 survived, 0 did not
- `name` - the passenger's name
- `sex` - the passenger's sex
- `age` - the passenger's age
- `sibsbp` - number of siblings on board
- `parch` - number of parents on board
- `ticket` - ticket number
- `fare` - how much the passenger paid for their ticket in British Pounds
- `cabin` - the cabin number
- `embarked` - the location they embarked - C=Cherbourg, Q=Queenstown, S=Southampton
- `boat` - the boat number 
- `body` - body identification number
- `home.dest` - their home followed by the destination

You may notice that a lot of values show up as `NaN`. This could mean that the values were empty or just weren't capable of being read by pandas.
We can deal with these values in a variety of ways. We can eliminate them from our dataset or place guesstimates instead. Depending on how we will be using the data, each approach might have some value.
We will focus on eliminating the values since that is usually the better/safer approach.

If a particular feature doesn't have that many rows of data then we might want to eliminate that feature, otherwise we can delete the rows that have `NaN`s in them. First, let's find out more about our data. We can do so using the [describe](https://pandas.pydata.org/pandas-docs/stable/generated/pandas.DataFrame.describe.html) method of DataFrames.

```python
df.describe(include="all")
```

We can see from the above summary that the cabin and boat columns have the fewest values.

## Data Cleaning

We can also check how many values in each column are null. A good way to do this is explained in this [stack overflow question](https://stackoverflow.com/a/26266451/818687). Find out how many `NaN`s we have in each column.

<!-- grade_id: cell-100bb17e0273b8d0, points: 10, grade: true, solution: true -->
```python
# Show the number of NaNs we have in each column
# YOUR CODE HERE
raise NotImplementedError()
```

Since some features have a lot of null values, we may want to drop those features instead of deleting all the rows that don't have values for them. We can do so using the [drop](https://pandas.pydata.org/pandas-docs/stable/generated/pandas.DataFrame.drop.html) method of DataFrames. We won't be revisiting these columns so you can set `inplace=True` when you use this method, this will update the original DataFrame. Drop the cabin, boat, body, and home.dest columns.

<!-- grade_id: cell-560efd89fc3e08b9, points: 0, grade: false, solution: true -->
```python
# Drop the cabin, boat, body, and home.dest columns.
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-cdc57e0b1c3d5fc7, points: 10, grade: true, solution: false -->
```python
assert "cabin" not in df.columns
assert "boat" not in df.columns
assert "body" not in df.columns
assert "home.dest" not in df.columns
```

```python
df.describe()
```

We still had some rows that had NaN values so let's drop those rows. We can use the [dropna](https://pandas.pydata.org/pandas-docs/stable/generated/pandas.DataFrame.dropna.html) method. Drop all rows that have any `NaN` values.
You can also just update the df variable since we won't be using the other data (set `inplace=True`).

<!-- grade_id: cell-3995d8b30799a584, points: 0, grade: false, solution: true -->
```python
# Drop any rows with na values. Use inplace=True or assign the result to df
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d521d0c8b45116d0, points: 10, grade: true, solution: false -->
```python
assert df.isnull().sum().sum() == 0
assert len(df) == 1043
```

```python
df.describe()
```

We know that some of our features are represented as numbers but are actually classes, such as `pclass` and `survived`. Let's look at how many people survived vs didn't and how many were in each class.

We can use the [value_counts](https://pandas.pydata.org/pandas-docs/stable/generated/pandas.Series.value_counts.html) Series method for this. First, let's look at how we select a feature column from a DataFrame. We can select it by passing in the name of the column just like we would with Python dictionaries.

```python
# Selecting the "survived" column
df["survived"]
```

When we select an individual DataFrame column it returns the data as a Series object. You can read more about [Series](https://pandas.pydata.org/pandas-docs/stable/generated/pandas.Series.html) and check out the [user guide](https://pandas.pydata.org/pandas-docs/stable/dsintro.html#series).

Many of the methods available in pandas work for both DataFrames and Series. `head` and `tail` are examples of those.

```python
df["survived"].head()
```

```python
df["survived"].tail()
```

Now let's try out [value_counts](https://pandas.pydata.org/pandas-docs/stable/generated/pandas.Series.value_counts.html) to determine how many people survived and how many died.

```python
df["survived"].value_counts()
```

We can also determine these values as a percentage by setting `normalize=True`.

```python
df["survived"].value_counts(normalize=True)
```

Now determine the same for the passenger classes and the sexes.

<!-- grade_id: pclass-value_counts, points: 5, grade: true, solution: true -->
```python
# Find out how many people were in each class
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: pclass-value_counts_norm, points: 5, grade: true, solution: true -->
```python
# Find out what percentage of people were in each class
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: sex-value_counts, points: 5, grade: true, solution: true -->
```python
# Find out how many people are of each sex
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: sex_value_counts_norm, points: 5, grade: true, solution: true -->
```python
# Find out what percentage of people are of each sex
# YOUR CODE HERE
raise NotImplementedError()
```

Value counts can give us some useful insights into the data and it works with string values and with numeric values as well. When we start working with machine learning models however, they prefer that things are in numeric form and not strings. If we also have multiple classes of a certain value it's better to have them as multiple features with each one specifying whether or not the value is true. These true/false features for each class are called dummy variables. Pandas makes it easy to create these dummy variables using the [get_dummies](https://pandas.pydata.org/pandas-docs/stable/generated/pandas.get_dummies.html) function. Try it out on the `pclass` and `sex` columns.

<!-- grade_id: cell-a654a8408c69c6e9, points: 5, grade: true, solution: true -->
```python
# Get the dummy variables for the `sex` column 
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-612600388764bb32, points: 5, grade: true, solution: true -->
```python
# Get the dummy variables for the `pclass` column 
# YOUR CODE HERE
raise NotImplementedError()
```

Dummy variables will be extremely useful when we start applying machine learning. For now, let's go back to data selection.

## Data Selection

We discussed how to select one column of a DataFrame, how that is just like we would use a dictonairy and how it returns a Series object. What about if we want to select multiple columns?

The notation for selecting multiple columns is to pass a list of strings instead of just one string. In the remaining section of this case study we will only use the `survived`, `age`, `sex`, and `pclass` features. Select those columns in a variable called `smallDf`.

Since we will be editing this variable, make sure you create a [copy](https://pandas.pydata.org/pandas-docs/stable/generated/pandas.DataFrame.copy.html) after selecting the columns. This [stack overflow answer](https://stackoverflow.com/a/27680109/818687) explains why we want to create a copy quite well.

<!-- grade_id: cell-dd971fee5d6aa470, points: 0, grade: false, solution: true -->
```python
# Select the survived, age, sex and pclass columns of the df and save them as smallDf
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3e9e5d1bf13dcf26, points: 5, grade: true, solution: false -->
```python
assert "survived" in smallDf.columns
assert "age" in smallDf.columns
assert "sex" in smallDf.columns
assert "pclass" in smallDf.columns
assert len(smallDf) == 1043
```

```python
smallDf.head()
```

Indexing in pandas can be very powerful. The [guide on indexing](https://pandas.pydata.org/pandas-docs/stable/indexing.html) is extremely useful. We'll cover only a small portion of what you can do with indexing, namely boolean indexing. 

First, we can apply a logical operation to a Series to return a new Series that checks whether the logical operation is true. 
Second, we can use that result to only select a small portion of the data.

This might sound complicated but hopefully an example can show how to use Boolean Indexing.

```python
# Create a Series that lets us know whether a person is male
smallDf["sex"] == "female"
```

```python
# Select all males from the dataset
malesDf = smallDf[smallDf["sex"] == "male"]
len(malesDf)
```

```python
malesDf.describe(include="all")
# Notice how there are no females in this df
```

<!-- grade_id: cell-1225f02668d8b333, points: 0, grade: false, solution: true -->
```python
# Now create a femalesDf which has data about the females only
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-faf20230f201c749, points: 5, grade: true, solution: false -->
```python
assert all(femalesDf["sex"] == "female")
assert len(femalesDf) == 386
```

This approach may be useful if we wanted to only select one portion of data that meets a certain logical structure. This example showed how we can apply it to string values but we can also apply this to numeric values. Let's say we wanted to determine how many people in our dataset are under the age of 10. We can determine that using the following:

```python
underTen = smallDf["age"] < 10
underTen.sum()
```

We can also apply more complex functionality to our data using the apply method for [DataFrames](https://pandas.pydata.org/pandas-docs/stable/generated/pandas.DataFrame.apply.html) or [Series](https://pandas.pydata.org/pandas-docs/stable/generated/pandas.Series.apply.html).

We can assign new columns to a DataFrame just as we would assign a new value for a dictionary.
Use the apply method and create a new column called "ageGroup" that has a value of "child" for people under the age of 13, "teen" for ages 13-19, and "adult" for ages above 19.

<!-- grade_id: cell-69d4fbe08865f184, points: 10, grade: true, solution: true -->
```python
# Use the apply method and create a new column called "ageGroup" 
# It should have a value of "child" for people under the age of 13, 
# "teen" for ages 13-19, and "adult" for ages above 19.
# You may want to write a function here to use with apply
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-559ab8938dcb0320, points: 0, grade: true, solution: false -->
```python
assert "ageGroup" in smallDf.columns
```

For the male/female case, if we wanted to determine the number of each group we could actually use the `groupby` functionality instead. Grouping is also another very powerful feature of pandas for data manipulation and the [guide](https://pandas.pydata.org/pandas-docs/stable/groupby.html) is really helpful.

```python
smallDf.groupby(["survived", "pclass", "sex"]).count()
```

When we group by some values we can perform some useful calculations such as count, mean, sum and others. The order of the grouping determines the breakdown of the data.

Repeat the above grouping in another order that you are curious about and perform a calculation that you find interesting.

<!-- grade_id: cell-029328c41ddbaca6, points: 10, grade: true, solution: true -->
```python
# Repeat the above grouping in another order that you are curious about
# perform a calculation that you find interesting ex. mean, count, sum
# YOUR CODE HERE
raise NotImplementedError()
```

This concludes our pandas review. Please share your feedback below.

## Feedback

<!-- grade_id: cell-74727d848a4325b6, points: 0, grade: false, solution: true -->
```python
def feedback():
    """Provide feedback on the contents of this exercise
    
    Returns:
        string
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e07ad32acc7178ee, points: 0, grade: true, solution: false -->
```python

```