 As this is a case study individual assignment, I agree and acknowledge that all code modified in this notebook is my own. I have not and will not collaborate with anyone on this assignment. If I have questions, I will ask the instructor or TAs.

<!-- grade_id: cell-7a410150ea339a11, points: 0, grade: false, solution: true -->
```python
# Please provide your name and agreement to the above statement as variables `name` (string) and `agree` (boolean)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-be8f2fd5450b57d8, points: 0, grade: true, solution: false -->
```python
conditions = [isinstance(name, str), isinstance(agree, bool), agree]
for condition in conditions:
    if not condition:
        raise ValueError("Student has not agreed to work on this assignment alone and without collaboration or has not provided their name")
```

# Job Performance Prediction

You work for a software startup, Predict All The Things Inc. (PALT), and are approached by the CEO to build an algorithm that can help sift through resumes. PALT just closed a $3 million Series A round of funding and the CEO just landed a deal with a national retailer, SellsALOT, to help them with hiring Sales Associates.

They are able to obtain data on all the employees that work as Sales Associates throughout their stores as well as customer satisfaction and sales performance scores.

In this case study, you are tasked to build a model to predict job performance to assist HR in selecting applicants to interview.

The data was provided to you by the new HR intern, Keegan. This is the email you got from Keegan with the attached data.

>Hi!
>
>I hope you're doing well. I've attached the data we have about all employees. Please ensure this data stays confidential and is not shared with anyone who has not signed the NDA. The columns have all the information we have about our employees and the scoring rating that they've received from our performance monitors. We also have some employees that were fired and I have included those as well.
>
>I was also able to dig up some more information about our employees that I found on the internet. It took a lot of time but I hope it helps in making the model even better. Can't wait to see this thing in action. Everyone here is very excited about our collaboration with you and we look forward to this making hiring a lot easier for us.
>
>Thanks,
>
>Keegan Thiel
>
>HR Intern
>
>Human Resources
>
>SellsALOT


Data is available in the `employees.csv` file provided. 


SellsALOT is an Equal Opportunity Employer which is an employer who agrees not to discriminate against any employee or job applicant because of race, color, religion, national origin, sex, physical or mental disability, or age.


## Data Cleaning

First, let's investigate the data that we received from Keegan.



```python
import pandas as pd
import matplotlib
import numpy as np
import sklearn as sk
from sklearn.model_selection import train_test_split
import datetime
from datetime import date
```

```python
df = pd.read_csv("employees.csv")
```

```python
df.head()
```

```python
df.describe(include="all")
```

```python
print("The columns of data are:")
list(df.columns)
```

Based on the data above we need to do the following:

1. Split Myers Briggs into subtypes
1. Convert categorical columns to dummy variables
1. Calculate Age based on date of birth

The [Myers Briggs Type Indicator](https://en.wikipedia.org/wiki/Myers%E2%80%93Briggs_Type_Indicator) (MBTI) descibes people as one of two types for each of:

* extraversion (E) or introversion (I)
* sensing (S) or intuition (N)
* thinking (T) or feeling (F)
* judgment (J) or perception (P)

It would make more sense for us to represent people as one or the other of these instead of creating all the possible cases. That way a model can learn based on each of those factors as well as their combination. 

Your next task is to split the MBTI column into four columns in the dataframe:

* MBTI_EI with value `E` or `I`
* MBTI_SN with value `S` or `N`
* MBTI_TF with value `T` or `F`
* MBTI_JP with value `J` or `P`

That correspond to the same row's Myers Briggs Type.

<!-- grade_id: cell-b96d43f098c26707, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1b5cf11c2512def1, points: 10, grade: true, solution: false -->
```python
assert len(set(df["MBTI_EI"])) == 2
assert "E" in set(df["MBTI_EI"]) and "I" in set(df["MBTI_EI"])
assert len(set(df["MBTI_SN"])) == 2
assert "S" in set(df["MBTI_SN"]) and "N" in set(df["MBTI_SN"])
assert len(set(df["MBTI_TF"])) == 2
assert "T" in set(df["MBTI_TF"]) and "F" in set(df["MBTI_TF"])
assert len(set(df["MBTI_JP"])) == 2
assert "J" in set(df["MBTI_JP"]) and "P" in set(df["MBTI_JP"])
```

1. ~~Split Myers Briggs into subtypes~~
1. Convert categorical columns to dummy variables
1. Calculate Age based on date of birth

Dumy variables are variables that allow us to convert a category into several binary variables. For example, if we had a color value that we were storing and we knew it could only have the values `red`, `green`, and `blue`, then instead of storing the color as those strings, we can store three binary variables: `is_red`, `is_green`, and `is_blue`. 

We can do this in pandas easily by using [`get_dummies`](https://pandas.pydata.org/pandas-docs/stable/generated/pandas.get_dummies.html).

<!-- grade_id: cell-32271c292a515ef3, points: 0, grade: false, solution: true -->
```python
# Determine all the categorical columns and save them to categorical_columns
# Note that this does not include the binary features

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-528eea88864f2926, points: 10, grade: true, solution: false -->
```python
assert len(categorical_columns) > 8
for category in categorical_columns:
    assert category in df.columns
```

<!-- grade_id: cell-ad307669e4346527, points: 0, grade: false, solution: true -->
```python
# For every column in the categorical features
# Calculate the dummy variables and add them to the dataframe

# YOUR CODE HERE
raise NotImplementedError()
```

```python
assert len(list(df.columns)) > 45
```

```python
print("The current columns are:")
list(df.columns)
```

<!-- grade_id: cell-10890eae4243cb8d, points: 0, grade: false, solution: true -->
```python
# Now let's drop all the categorical features columns from the dataframe
# So that we don't have duplicate information stored
# YOUR CODE HERE
raise NotImplementedError()
```

```python
print("The current columns are:")
list(df.columns)
```

<!-- grade_id: cell-a6172310793c3289, points: 10, grade: true, solution: false -->
```python
assert 45 > len(list(df.columns)) > 30
```

1. ~~Split Myers Briggs into subtypes~~
1. ~~Convert categorical columns to dummy variables~~
1. Calculate Age based on date of birth

```python
def calculate_age(born):
    """Calculates age based on date of birth using https://stackoverflow.com/a/9754466/818687

    Args:
        born (datetime): The date of birth

    Returns:
        int: The age based on date of birth
    """
    today = date.today()
    return today.year - born.year - ((today.month, today.day) < (born.month, born.day))
```

<!-- grade_id: cell-315ed06d1f0428b4, points: 0, grade: false, solution: true -->
```python
# Add an "Age" column to the dataframe that calculates people's ages based on their date of birth
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-abfc46abea98fba7, points: 20, grade: true, solution: false -->
```python
assert df["Age"].min() == 20
assert df["Age"].max() == 36
assert df["Age"].median() == 28
```

## Modelling

Based on your understanding of the data, select the features that you want to use to predict:

1. Customer Satisfaction
1. Sales Performance
1. Fired

<!-- grade_id: cell-449eb7f2253cca83, points: 0, grade: false, solution: true -->
```python
# Save the columns we are trying to predict to targets
# Make sure that if we had a categorical column, that you use the dummy representation(s)
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-458aa3767005cc6d, points: 20, grade: true, solution: false -->
```python
assert len(targets) == 3
for target in targets:
    assert target in df.columns
```

Your prediction will be used to rank applicants for interviews with HR. **Which features will you select to use in your model?**

```python
print("The available columns are:")
list(df)
```

<!-- grade_id: cell-bd2399d20615cbd1, points: 0, grade: false, solution: true -->
```python
# Enter all the features you want to use in a list and save it to rank_features
# These are the features for the model that will rank applicants for interviews
# YOUR CODE HERE
raise NotImplementedError()
```

Why did you choose the features you did?

<!-- grade_id: cell-11725eab2a5dc436, points: 0, grade: false, solution: true -->
```python
## Save your reasoning in a string to the variable ranking_reason

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8b0e41f067c8ef32, points: 20, grade: true, solution: false -->
```python
assert isinstance(ranking_reason, str)
assert len(ranking_reason) > 20
```

<!-- grade_id: cell-d9d5cf91c533af77, points: 0, grade: false, solution: true -->
```python
# Perform a train and test split on the data with the variable names:
# rank_x_train for the training features
# rank_x_test for the testing features
# rank_y_train for the training targets
# rank_y_test for the testing targets
# The test dataset should be 20% of the total dataset

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-146c7b1df2f26d90, points: 10, grade: true, solution: false -->
```python
assert (len(rank_x_train) / (len(rank_x_test) + len(rank_x_train))) == 0.8
assert (len(rank_y_train) / (len(rank_y_test) + len(rank_y_train))) == 0.8
assert len(rank_x_train) == len(rank_y_train)
assert len(rank_x_test) == len(rank_y_test)
```

<!-- grade_id: cell-99742e48e538f29a, points: 50, grade: true, solution: true -->
```python
# Select models of your choosing, import them and perform a parameter search to train them on each of the targets
# Determine an appropriate metric for measuring your performance and report that.

# YOUR CODE HERE
raise NotImplementedError()

# You should have print messages that state something like:
# To predict the target (target), I trained a (model) model
# and determined the best hyperparameters as (param1 = p1), (param2 = p2)...
# resulting in a (metric) score of (score)
```

Would your feature choice change if HR was going to use your model to directly hire applicants without an interview? **Which features will you select to use in that model?**

```python
print("The available columns are:")
list(df)
```

<!-- grade_id: cell-c4bad209682c7022, points: 0, grade: false, solution: true -->
```python
# Enter all the features you want to use in a list and save it to selection_features
# These are the features for the model that will directly hire the top applicants
# YOUR CODE HERE
raise NotImplementedError()
```

Why did you choose the features you did?

<!-- grade_id: cell-b95231da655a979d, points: 0, grade: false, solution: true -->
```python
## Save your reasoning in a string to the variable selection_reason

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-34d2462a6b2f374c, points: 20, grade: true, solution: false -->
```python
assert isinstance(selection_reason, str)
assert len(selection_reason) > 20
```

Why was your choice different from or the same as the ranking features?


<!-- grade_id: cell-4479d690eeb487ad, points: 0, grade: false, solution: true -->
```python
# Save your reasoning in a string to the variable
# same_reason if the features are the same
# different_reason if the features are different
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f875ee0056d6eb98, points: 20, grade: true, solution: false -->
```python
if all([rf in selection_features for rf in rank_features]) and all([sf in rank_features for sf in selection_features]):
    print("Your features for ranking and selection are the same.")
    assert isinstance(same_reason, str)
    assert len(same_reason) > 20
else:
    print("Your features for ranking and selection are different.")
    assert isinstance(different_reason, str)
    assert len(different_reason) > 20
```

<!-- grade_id: cell-fe98928b7423b02b, points: 0, grade: false, solution: true -->
```python
# Perform a train and test split on the data with the variable names:
# selection_x_train for the training features
# selection_x_test for the testing features
# selection_y_train for the training targets
# selection_y_test for the testing targets
# The test dataset should be 20% of the total dataset

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-be29120b1249cd74, points: 10, grade: true, solution: false -->
```python
assert (len(selection_x_train) / (len(selection_x_test) + len(selection_x_train))) == 0.8
assert (len(selection_y_train) / (len(selection_y_test) + len(selection_y_train))) == 0.8
assert len(selection_x_train) == len(selection_y_train)
assert len(selection_x_test) == len(selection_y_test)
```

Now let's see if the model performs differently.

<!-- grade_id: cell-43f1427eb940e017, points: 30, grade: true, solution: true -->
```python
# Select models of your choosing, import them and perform a parameter search to train them on each of the targets
# Determine an appropriate metric for measuring your performance and report that.

# YOUR CODE HERE
raise NotImplementedError()

# You should have print messages that state something like:
# To predict the target (target), I trained a (model) model
# and determined the best hyperparameters as (param1 = p1), (param2 = p2)...
# resulting in a (metric) score of (score)
```

<!-- grade_id: cell-5df6d98d3b874670, points: 20, grade: true, solution: true -->
```python
# Follow this up with a comparison between the performance on your 2 models using the different features.
# You should print something like
# Using rank features for target (target) the model scored (score)
# versus using the selection features where it scored (score)
# YOUR CODE HERE
raise NotImplementedError()

```

## Feedback

<!-- grade_id: cell-259e62fb9f936804, points: 0, grade: false, solution: true -->
```python
def feedback():
    """Provide feedback on the contents of this exercise
    
    Returns:
        string
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-865e5f961cfd020f, points: 0, grade: true, solution: false -->
```python

```